# Assignment 1: Python Skills

## Instructions

In this assignment, you will answer some conceptual questions about Python and write a collection of basic algorithms and data structures.  You will find that in addition to a problem specification, each programming question also includes examples from the Python interpreter. These are meant to illustrate typical use cases, and should not be taken as comprehensive test suites.

### Note:
You are expected to write code where you see **your code here**.  
Make sure you delete the lines with **pass** and **raise NotImplementedError** or your code may not run correctly.

## Working with Lists

Write a function concatenate(seqs) that returns a list containing the concatenation of the elements of the input sequences. Your implementation should consist of a single list comprehension, and should not exceed one line.

<!-- grade_id: cell-solution-1, points: 0, grade: false, solution: true -->
```python
############################################################
# Section: Working with Lists
############################################################

def concatenate(seqs):
    return [item for seq in seqs for item in seq]

```

<!-- grade_id: cell-test-1, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

# concatenate test
concatenate_test = concatenate([[1, 2], [3, 4]])
assert concatenate_test == [1, 2, 3, 4], "Concatenated list should be [1, 2, 3, 4]"
concatenate_test = concatenate([[1, 2], [3, 4, 5]])
assert concatenate_test == [1, 2, 3, 4, 5], "Concatenated list should be [1, 2, 3, 4, 5]"
print("test passed!")

```

Output:
```
test passed!
```

Write a function transpose(matrix) that returns the transpose of the input matrix, which is represented as a list of lists. Recall that the transpose of a matrix is obtained by swapping its rows with its columns. More concretely, the equality matrix[i][j] == transpose(matrix)[j][i] should hold for all valid indices i and j. You may assume that the input matrix is well-formed, i.e., that each row is of equal length. You may further assume that the input matrix is non-empty. Your function should not modify the input.

<!-- grade_id: cell-solution-2, points: 0, grade: false, solution: true -->
```python
def transpose(matrix):
    return [[row[i] for row in matrix] for i in range(len(matrix[0]))]

```

<!-- grade_id: cell-test-2, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

# transpose test
transpose_test1 = transpose([[1, 2, 3]])
transpose_test2 = transpose([[1, 2], [3, 4], [5, 6]])
assert transpose_test1 == [[1], [2], [3]], "Transposed list should be [[1], [2], [3]]"
assert transpose_test2 == [[1, 3, 5], [2, 4, 6]], "Transposed list should be [[1, 3, 5], [2, 4, 6]]"
print("test passed!")

```

Output:
```
test passed!
```

## Sequence Slicing

Write a function copy(seq) that returns a new sequence containing the same elements as the input sequence.

<!-- grade_id: cell-solution-3, points: 0, grade: false, solution: true -->
```python
############################################################
# Section: Sequence Slicing
############################################################

def copy(seq):
     return seq[:]

```

<!-- grade_id: cell-test-3, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

# copy test
copy_test1 = copy("abc")
copy_test2 = copy((1, 2, 3))
assert copy_test1 == 'abc'
assert copy_test2 == (1, 2, 3)
print("test passed!")

```

Output:
```
test passed!
```

Write a function all_but_last(seq) that returns a new squence containing all but the last element of the input sequence. If the input sequence is empty, a new empty sequence of the same type should be returned.

<!-- grade_id: cell-solution-4, points: 0, grade: false, solution: true -->
```python
def all_but_last(seq):
     return seq[:-1]

```

<!-- grade_id: cell-test-4, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

# all_but_last test
all_but_last_test1 = all_but_last("abc")
all_but_last_test2 = all_but_last((1, 2, 3))
assert all_but_last_test1 == 'ab'
assert all_but_last_test2 == (1, 2)
print("test passed!")

```

Output:
```
test passed!
```

Write a function every_other(seq) that returns a new sequence containing every other element of the input sequence, starting with the first. This function can be written in one line using the optional third parameter of the slice notation.

<!-- grade_id: cell-solution-5, points: 0, grade: false, solution: true -->
```python
def every_other(seq):
    return seq[::2]

# your code here

```

<!-- grade_id: cell-test-5, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

# every_other test
every_other_test1 = every_other([1, 2, 3, 4, 5])
every_other_test2 = every_other("abcde")
assert every_other_test1 == [1, 3, 5]
assert every_other_test2 == 'ace'
print("test passed!")

```

Output:
```
test passed!
```

## Combinatorial Algorithms

The prefixes of a sequence include the empty sequence, the first element, the first two elements, etc., up to and including the full sequence itself.  

Write a function prefixes(seq) that yield all prefixes of the input sequence.

<!-- grade_id: cell-solution-6, points: 0, grade: false, solution: true -->
```python
############################################################
# Section: Combinatorial Algorithms
############################################################

def prefixes(seq):
    for i in range(len(seq) + 1):
        yield seq[:i]

```

<!-- grade_id: cell-test-6, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

# prefixes test
prefixes_test1 = list(prefixes([1, 2, 3]))
prefixes_test2 = list(prefixes("abc"))
assert prefixes_test1 == [[], [1], [1, 2], [1, 2, 3]]
assert prefixes_test2 == ['', 'a', 'ab', 'abc']
print("test passed!")

```

Output:
```
test passed!
```

Write a function slices(seq) that yields all non-empty slices of the input sequence.

<!-- grade_id: cell-solution-7, points: 0, grade: false, solution: true -->
```python
def slices(seq):
    for start in range(len(seq)):
        for end in range(start + 1, len(seq) + 1):
            yield seq[start:end]


```

<!-- grade_id: cell-test-7, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

# slices test
slices_test1 = list(slices([1, 2, 3]))
slices_test2 = list(slices("abc")) 
assert slices_test1 == [[1], [1, 2], [1, 2, 3], [2], [2, 3], [3]]
assert slices_test2 == ['a', 'ab', 'abc', 'b', 'bc', 'c']
print("test passed!")

```

Output:
```
test passed!
```

## Text Processing

A common preprocessing step in many natural language processing tasks is text normalization, wherein words are converted to lowercase, extraneous whitespace is removed, etc. 

Write a function normalize(text) that returns a normalized version of the input string, in which all words have been converted to lowercase and are separated by a single space. No leading or trailing whitespace should be present in the output.

<!-- grade_id: cell-solution-8, points: 0, grade: false, solution: true -->
```python
############################################################
# Section: Text Processing
############################################################

def normalize(text):
    return ' '.join(text.lower().split())

```

<!-- grade_id: cell-test-8, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

# normalize test
normalize_test1 = normalize("This is an example.")
normalize_test2 = normalize(" EXTRA SPACE ")
assert normalize_test1 == "this is an example."
assert normalize_test2 == "extra space"
print("test passed!")

```

Output:
```
test passed!
```

Write a function digits_to_words(text) that extracts all digits from the input string, spells them out as lowercase English words, and returns a new string in which they are each separated by a single space. If the input string contains no digits, then an empty string should be returned.

<!-- grade_id: cell-solution-9, points: 0, grade: false, solution: true -->
```python
def digits_to_words(text):
    digit_to_word = {'0': 'zero', '1': 'one', '2': 'two', '3': 'three', '4': 'four',
                     '5': 'five', '6': 'six', '7': 'seven', '8': 'eight', '9': 'nine'}
    
    digits = ''.join(filter(str.isdigit, text))
    
    return ' '.join(digit_to_word[d] for d in digits) if digits else ''

```

<!-- grade_id: cell-test-9, points: 10, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

# digits_to_words test
digits_to_words_test1 = digits_to_words("Zip Code: 19104")
digits_to_words_test2 = digits_to_words("Pi is 3.1415...")
assert digits_to_words_test1 == "one nine one zero four"
assert digits_to_words_test2 == "three one four one five"
print("test passed!")

```

Output:
```
test passed!
```

<!-- grade_id: cell-solution-10, points: 0, grade: false, solution: true -->
```python
# your code here
print("testing over")
```

Output:
```
testing over
```

## Python Packages

Install numpy first:

<!-- grade_id: cell-solution-11, points: 0, grade: false, solution: true -->
```python
# your code here
! pip install numpy
```

Output:
```
[33mWARNING: The directory '/home/jovyan/.cache/pip' or its parent directory is not owned or is not writable by the current user. The cache has been disabled. Check the permissions and owner of that directory. If executing pip with sudo, you should use sudo's -H flag.[0m[33m
[0mRequirement already satisfied: numpy in /opt/conda/lib/python3.10/site-packages (1.23.5)
[33mWARNING: Running pip as the 'root' user can result in broken permissions and conflicting behaviour with the system package manager. It is recommended to use a virtual environment instead: https://pip.pypa.io/warnings/venv[0m[33m
[0m
[1m[[0m[34;49mnotice[0m[1;39;49m][0m[39;49m A new release of pip is available: [0m[31;49m24.0[0m[39;49m -> [0m[32;49m24.1.2[0m
[1m[[0m[34;49mnotice[0m[1;39;49m][0m[39;49m To update, run: [0m[32;49mpip install --upgrade pip[0m
```


Use numpy package to implement the sort_array(list_of_matrics) which takes in a list of matrices with various dimensions and return a sorted 1D array with decreasing order contains all the values in these matrices. The data type of the returned array is int.

<!-- grade_id: cell-solution-13, points: 0, grade: false, solution: true -->
```python
############################################################
# Section: Python Packages
############################################################
import numpy as np
def sort_array(list_of_matrices):
    flattened_arrays = [matrix.flatten() for matrix in list_of_matrices]
    concatenated_array = np.concatenate(flattened_arrays)
    sorted_array = np.sort(concatenated_array)[::-1]  # Sort in decreasing order
    return sorted_array.astype(int)

```

<!-- grade_id: cell-test-13, points: 5, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

# sort_array test
from numpy.testing import assert_array_equal
matrix1 = np.array([[1, 2], [3, 4]])
matrix2 = np.array([[5, 6, 7], [7, 8, 9], [0, -1, -2]])
sort_array_test = sort_array([matrix1, matrix2])

expected_result = np.array([9, 8, 7, 7, 6, 5, 4, 3, 2, 1, 0, -1, -2])
assert_array_equal(sort_array_test, expected_result)

print("test passed!")


```

Output:
```
test passed!
```

<!-- grade_id: cell-solution-14, points: 0, grade: false, solution: true -->
```python
# your code here
print("numpy test over")
```

Output:
```
numpy test over
```

<!-- grade_id: cell-solution-15, points: 0, grade: false, solution: true -->
```python
# your code here
! pip install nltk
```

Output:
```
[33mWARNING: The directory '/home/jovyan/.cache/pip' or its parent directory is not owned or is not writable by the current user. The cache has been disabled. Check the permissions and owner of that directory. If executing pip with sudo, you should use sudo's -H flag.[0m[33m
[0mCollecting nltk
  Downloading nltk-3.8.1-py3-none-any.whl.metadata (2.8 kB)
Requirement already satisfied: click in /opt/conda/lib/python3.10/site-packages (from nltk) (8.1.3)
Requirement already satisfied: joblib in /opt/conda/lib/python3.10/site-packages (from nltk) (1.2.0)
Collecting regex>=2021.8.3 (from nltk)
  Downloading regex-2024.5.15-cp310-cp310-manylinux_2_17_x86_64.manylinux2014_x86_64.whl.metadata (40 kB)
[2K     [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m40.9/40.9 kB[0m [31m6.4 MB/s[0m eta [36m0:00:00[0m
[?25hRequirement already satisfied: tqdm in /opt/conda/lib/python3.10/site-packages (from nltk) (4.64.1)
Downloading nltk-3.8.1-py3-none-any.whl (1.5 MB)
[2K   [90m━━━━━━━━
```


Implement the POS_tag(sentence) function which takes in a string and return the Part-of-Speech(POS) tag of each word in the sentence. To complete this task, you need to fulfill the requirements shown below:
1) Convert all the characters in the sentence to lower case. 
2) Tokenize the sentence.
3) Remove the stop words and punctuation.
4) Conduct the pos tagging and return a list of tuples.


<!-- grade_id: cell-solution-16, points: 0, grade: false, solution: true -->
```python
import string
import nltk
from nltk.corpus import stopwords
from nltk.tokenize import word_tokenize
# Download necessary NLTK data files (only need to run this once)
nltk.download('punkt')
nltk.download('averaged_perceptron_tagger')
nltk.download('stopwords')
def POS_tag(sentence):
    # Convert sentence to lowercase
    sentence = sentence.lower()
    
    # Tokenize the sentence
    tokens = word_tokenize(sentence)
    
    # Remove punctuation
    tokens = [token for token in tokens if token.isalnum()]
    
    # Define stop words
    stop_words = set(stopwords.words('english'))
    
    # Remove stop words
    tokens = [token for token in tokens if token not in stop_words]
    
    # Perform part-of-speech tagging
    tagged_tokens = nltk.pos_tag(tokens)
    
    return tagged_tokens
```

Output:
```
[nltk_data] Downloading package punkt to /home/jovyan/nltk_data...
[nltk_data]   Package punkt is already up-to-date!
[nltk_data] Downloading package averaged_perceptron_tagger to
[nltk_data]     /home/jovyan/nltk_data...
[nltk_data]   Package averaged_perceptron_tagger is already up-to-
[nltk_data]       date!
[nltk_data] Downloading package stopwords to /home/jovyan/nltk_data...
[nltk_data]   Package stopwords is already up-to-date!
```

<!-- grade_id: cell-test-16, points: 5, grade: true, solution: false -->
```python
##########################
### TEST YOUR SOLUTION ###
##########################

# POS_tag test
sentence = 'The Force will be with you. Always.'
POS_tag_test = POS_tag(sentence)
assert POS_tag_test == [('force', 'NN'), ('always', 'RB')]
print("test passed!")

```

Output:
```
test passed!
```

```python
print("All testings over and all tests passed-check for all cases and run again if needed")
```

Output:
```
All testings over and all tests passed-check for all cases and run again if needed
```

```python

```